```python
%%capture
# OTTER IGNORE
! pip install drexel-jupyter-logger
! pip install otter-grader
! pip install pytest
! pip install scipy
```

```python
# Initialize Otter
import otter
grader = otter.Notebook("ENGR131_Final_W23.ipynb")
```

# ENGR 131

## Final Exam

## 10:30am - 12:30 pm

## Winter 2023

## Profs. Capps and Agar

**Instructions:** Make sure to read all of the instructions they will provide you with important information to help you through the exam. All the exam makes use of autograder. We will use the same tests to grade the exam. Make sure what you turn in earns as many points as possible using these tests. 

You should add code everywhere there is a `...`

**Rules:**

1. You must record your screen through the entirety of the exam. The video recording of your exam should be uploaded [here](https://drexel0-my.sharepoint.com/:f:/g/personal/sc3623_drexel_edu/Eo30-GfMK0BMr7GDffYZdjAB5GSyjDzAag6Pt12CzW-2vg).
2. The exam is closed book. You are not allowed to use anything outside of the jupyterhub. You are not allowed to search on the web, use your phone or tablet device, etc.
3. Your screen must be maximized through the exam and your clock on your computer must be visible.
4. You **may** and should look at docstring within JupyterHub.
5. If you encounter any problems with the recording, please let the instructor know as soon as possible.
6. If your laptop's power reduces to 25 percent, please notify the instructor. It is recommended that you run your computer in power saver mode if you have it. 
7. You should only end your recording after you turn in your assignment and log file on Drexel Learn.


If you would find it helpful to have a side-by-side view of the problem and instructions this is possible by `right-clicking` on the tab and selecting `new view for notebook`

![](https://github.com/DrexelEngineering/ENGR131_W2023/blob/main/jupyterbook/final/assets/side_by_side.png?raw=true)

**GOOD LUCK**


_Type your answer here, replacing this text._

```python
import drexel_jupyter_logger

# type your first and last name as a string
name_ = ...

```

```python
grader.check("q0-Type Your Name")
```

## Question 1: Making a Graph (12 points)

A friend gives you a card with the following function written on it. You need to figure out what they are saying. Implement a function to reveal the hidden message.

1. Import the packages you need.

   - You will need `numpy` import this as the variable `np`
   - You will need the `pyplot` module from `matplotlib`. Import `pyplot` as a variable `plt`.


2. Make a function called `draw_card_eq`.
   - This should take one input `a`.
   - Create a linearly-spaced vector using the `linspace` method in `np`.
     - The vector should go from 0 to $2\pi$ with 2000 steps.
     - Assign this vector to the variable `t`.
   - Write the equation for `x`:
     $$ x = a(16(\sin(t))^3)$$
   - Write the equation for `y`:
     $$ y = a(13\cos(t) - 5\cos(2t) - 2\cos(3t) - \cos(4t))$$

   - Use the `plot` method in `plt` to plot `x` vs `y`, set the plot color to `red` using the input `red` as a string
     - Save this to the variable `plot_`.
   - Return `plot_`.

3. Call the function `draw_card_eq` with `a` equal to 1. Doing so will reveal the message. 
   - Save the output of the function to a variable `plot_`.


```python
# Your import statements go here.
# We would recommend checking this with the autograder.
# Completion of this step should pass tests 1-2.
...

# Your function goes here
# We would recommend checking this with the autograder.
# Completion of this step should pass tests 3.
...

# Call your function here
# We would recommend checking this with the autograder.
# Completion of this step pass all tests.
...
```

```python
grader.check("q1 - Making a Graph")
```

## Question 2: Using Data Types (10 Points)

We have written a function that converts a DataType. Pass values with the expected DataTypes as defined below to the `convert_to_string` function. The function will print the output. 

0. Run the `convert_to_string` function first.
1. Pass an integer `42` to `convert_to_string`. 
    - Save the output to a variable `int_`.
2. Pass an float `3.14159` to `convert_to_string`. 
    - Save the output to a variable `float_`.
3. Pass an boolean with a true value to `convert_to_string`. 
    - Save the output to a variable `bool_`.
4. Pass an string `I love coding` to `convert_to_string`. 
    - Save the output to a variable `string_`.
5. Pass a list of values `1`, `2`, `3` to `convert_to_string`. 
    - Save the output to a variable `list_`.
6. Pass a tuple of values `a`, `b`, `c` to `convert_to_string`. 
    - Save the output to a variable `tuple_`.
7. Pass a dictionary with a key `name` and a value `jay` as a string and a key `age` and a value `12` as an integer to `convert_to_string`. 
    - Save the output to a variable `dict_`.



```python
def convert_to_string(data):
    if isinstance(data, int):
        print(data)
    elif isinstance(data, float):
        print("{:.2f}".format(data))
    elif isinstance(data, bool):
        print(str(data))
    elif isinstance(data, str):
        print(data)
    elif isinstance(data, list):
        elements = ", ".join([str(x) for x in data])
        print(elements)
    elif isinstance(data, tuple):
        elements = ", ".join([str(x) for x in data])
        print(elements)
    elif isinstance(data, dict):
        pairs = ", ".join(
            ["Key = " + str(k) + " Value = " + str(v) for k, v in data.items()]
        )
        print(pairs)
    return data
```

```python
# 1. Pass an integer `42` to `convert_to_string`. 
# Save the output to a variable `int_`.
# completion of this step will pass test 1
...
```

```python
# 2. Pass an float `3.14159` to `convert_to_string`. 
# Save the output to a variable `float_`.
# completion of this step will pass test 2
...
```

```python
# 3. Pass an boolean with a true value to `convert_to_string`. 
# Save the output to a variable `bool_`.
# completion of this step will pass test 3
...
```

```python
# 4. Pass an string `I love coding` to `convert_to_string`. 
# Save the output to a variable `string_`.
# completion of this step will pass test 4
...
```

```python
# 5. Pass a list of values `1`, `2`, `3` to `convert_to_string`. 
# Save the output to a variable `list_`.
# completion of this step will pass test 5
...
```

```python
# 6. Pass a tuple of values `a`, `b`, `c` to `convert_to_string`. 
# Save the output to a variable `tuple_`.
# completion of this step will pass test 6
...
```

```python
# 7. Pass a dictionary with a key `name` and a value `jay` as a string and a key `age` and a value `12` as an integer to `convert_to_string`. 
# Save the output to a variable `dict_`.
# completion of this step will pass test 7
...
```

```python
grader.check("q2 - Exploring DataTypes")
```

## Question 3: Building a Thermometer (16 Points)

In chemical engineering, it is common to monitor the temperature of a reactor. We will build a class to implement a temperature sensor. This sensor will save and also set a temperature value in Celsius. Additionally, the class will convert the temperature given in Celsius to Fahrenheit.

1. Build a class called `Thermometer`. This class will have to be initialized and have two methods.
    - Add an initialization method that accepts an initial temperature in Celsius as a variable `temperature`.
      - Save the `temperature` as an object attribute `temperature_c`. This is the temperature in $^{\circ}C$ .

    - Build a method `set_temperature` that will set the object attribute `temperature_c`.
      - This method takes one input `temperature`.
      - We want to make sure the input is a number. Do this by **trying** to convert it to a `float` using the built in `float` method.
        - If `temperature` is a number that can be converted to a float, save it to the object attribute `temperature_c`.
       - If `temperature` cannot be converted to a float, set the object attribute `temperature_c` to the string 'fail'.

    - Build a method `to_fahrenheit` that converts the temperature from Celsius to Fahrenheit.
      - Convert the temperature using the following equation.
        $$Fahrenheit = \frac{9}{5}Celsius +32 $$
      - Return the temperature in Fahrenheit.  

2. Instantiate a `Thermometer` `water_sensor` with an initial `temperature` of `12` $^{\circ}C$.
   
3. Extract the temperature in $^{\circ}C$ from the object `water_sensor`.
    - Assign the output to a variable `temp_c`.

4. Extract the temperature in $^{\circ}F$ from the object `water_sensor`.
    - Assign the output to a variable `temp_f`.


```python
# Your class for Thermometer goes here
...

# Instantiate the Thermometer water_sensor here
...

# extract the temperature in degrees celsius
...

# extract the temperature in degrees celsius
...
```

```python
grader.check("q3-Building a Thermometer")
```

## Question 4: Determining a Price for a Circuit (15 Points)

As an engineer, you might want to estimate the total price at which you would sell a circuit you designed based on the cost of goods from a supplier and a desired profit margin. We will write a Python function to assist you in this task.

We have provided you with two dictionaries, `parts_cost` and `circuit_parts`. `parts_cost` defines the the parts from the supplier, and `circuit_parts` defines the components needed to build the circuit. The circuit uses two `Custom Processor` components, which are not available from the supplier, that each costs $12.

*Cost of Goods from Supplier*

| Part Name (Key)     | Cost per Item (value, in $) |
| ---------------- | -------------- |
| 10-ohm resistor  | 0.025          |
| 100-ohm resistor | 0.03           |
| 1-pF capacitor   | 0.05           |
| 10-pF capacitor  | 0.12           |

*Components Required for Building the Circuit*

| Part Name  (Key)       | Quantity (sub-key) | Cost to Build (value, in $)|
| ---------------- | -------- | -------------- |
| 10-ohm resistor  | 10       | N/A             |
| 10-pF capacitor  | 3        | N/A            |
| 1-pF capacitor   | 7        | N/A            |
| Custom Processor | 2        | 12            |

1. Import the package `numpy` as `np`. 
    - You might not need to do this step if you have done it as part of the previous questions; however, it is good practice.
2. Assign `0.63` to a variable `profit`, which is the fractional profit you want to make when you sell this circuit. 
3. Define a function `print_total` to print the total price.
    - Return the string "Total cost of the order: \$\<sales_price\>", where \$\<sales_price\> is the value of the computed `sales_price`. 
    - For example, if `sales_price=3.00`, the string should read "Total cost of the order: $3.00". 
3. Define a function `price_calculator` that accepts three inputs: `parts_cost`, `circuit_parts`, and `profit`.
    - Initialize a variable `cost_of_goods`, and set it equal to 0.
    - Using a loop, compute the total `cost_of_goods`. 
        - You should use the `.items()` method of a dictionary that returns an iterator, where each iteration is a tuple containing the key and the value. 
        - You will need to use a branching statement to get the cost per item from the dictionary with the cost per item or the nested dictionary with the cost to build, whichever is appropriate for the part.
    - Determine the `sales_price` by adding the profit to the `cost_of_goods`. 
        - Multiply the `cost_of_goods` by (1 + `profit`). 
        - Using the `np.round` method, round the price to 2 decimal places. 
        - Save this value to the variable `sales_price`.
    - Call the `print_total` function within `price_calculator`.
        - Save the output to the variable `order_total`.
    - Print the variable `order_total` within the `price_calculator` function.
    - Return the variable `sales_price`.
8. Call the function `price_calculator` with the inputs `parts_cost`, `circuit_parts`, `profit`. 
    - Save the output to the variable `total_price`.


```python
# Your import statement goes here.
...

parts_cost = {
    "10-ohm resistor": 0.025,
    "100-ohm resistor": 0.03,
    "1-pF capacitor": 0.05,
    "10-pF capacitor": 0.12,
}

circuit_parts = {
    "10-ohm resistor": {"quantity": 10},
    "10-pF capacitor": {"quantity": 3},
    "1-pF capacitor": {"quantity": 7},
    "Custom Processor": {"quantity": 2, "price": 12},
}

# profit assignment goes here
# there are no tests associated with just this step.
...

# Your function for print total goes here
# There are no specific test associated with this step
...

# Your function price_calculator goes here
# Completion of this step should pass test 3
...

# Your call of the function price_calculator goes here
# Completion of this step should pass all tests
...
```

```python
grader.check("q4-Determining a Price for a Circuit")
```

## Plotting and Fitting (20 Points)

One of the most common tasks in engineering is fitting data to a model. For this question we will generate some noisy data with a gaussian distribution - one of the most common statistical distributions. We will then plot this data, and fit this data to the model.

1. Start by importing the necessary tools. If you imported any of these packages before, you might not need to do so again; however, it is good practice.

   - Import `numpy` as the variable `np`.
   - Import the `pyplot` module from `matplotlib` as a variable `plt`. 
   - Import the function `curve_fit` from `scipy.optimize`.

2. Define a function `gaussian` that computes the Gaussian function using the following equation:
   $$Ae^{\frac{-(x-b)^2}{2c^2}}$$
   - The function should accept inputs of `x`, `a`, `b`, and `c`.

3. Use the `linspace` method in `numpy` to generate a vector from `-20` to `20` with `100` steps. 
   - Save this vector to the variable `x_data`.

4. Use the defined `gaussian` function to compute the true, resulting value when `a` is `3`, `b` is `4`, and `c` is `7`. 
   - Save this to the variable `y_true`.

5. Generate noisy data using the `random.normal` method in `numpy`. 
   - The noisy data should be the same size as the `x_data`. 
   - Scale the noise by a factor of `0.3`, and add the noise the `y_true` variable.
   - Assign the output to `y_noisy`.

6. Using the `scatter` method from `matplotlib` `pyplot`, plot the `x_data` vs `y_noisy`. 
   - Add a label, "Noisy Data". 
   - Save this plot to the variable `plot_1`.

7. Use the imported `curve_fit` method to fit the noisy data to the `gaussian` function over the range of `x_data`. 
   - Save the fit results to the variable `popt`, and the parameter covariances to a variable `pcov`. 
     - **Do not provide any initial guess**, or provide `None`. If you provide an initial guess, the autograder might fail.

8. Calculate the fitted curve using the defined `gaussian` function. 
   - You can use the \* method to unpack the list of fitted parameters if you like. 
   - Save the result to a variable `y_fit`.

9.  Plot the fit results as `x_data` vs `y_fit` using the `plot` method in `pyplot`. 
    - Set the color of the line to green using the `g` tag.
    - Add a label, "Fitted Gaussian". 
    - Save this plot to the variable `plot_2`.

10. Using the methods within `plt`, set the `x_label` and `y_label` to `x` and `y`, respectively.
    
11. Use the `legend` method in `plt` to add a legend.


```python
# 1. Add your import statements here. 
# Completion of this step will pass the first test
...

# do not delete the line of code below, it will cause problems with the autograder.
np.random.seed(42)

# 2. Define the Gaussian function.
# completion of this step will pass test 2.
...

# 3. define your x_data vector
# completion of this step will pass test 3.
...

# 4. Call the gaussian function as described above
# save the output to the variable y_true
# completion of this step will pass test 4.
...

#5. Generate the noisy data described above
# save the output to the variable y_noisy
# completion of this step will pass test 5.
...

# 6. Plot the noisy data using a scatterplot
...

# 7. Fit the noisy data to a Gaussian function using curve_fit
# save the outputs to the variables popt and pcov
# completion of this step will pass test 6.
...

# 8. Calculate the fitted curve using the gaussian function
# save the output to the variable y_fit
# completion of this step will pass test 7. 
...

# 9. Plot the fitted results as described in the instructions
# save the output to the variable plot_2
# completion of this step will pass several more tests. 
...

# 10. Add axis labels to the plot
# completion of this step will pass a test.
...

# 11. Add a legend to the plot
# completion of this step will pass all tests.
...
```

```python
grader.check("q5-Plotting and Fitting")
```

Tests for q0-Type Your Name (points: None)
```
>>> import drexel_jupyter_logger
>>> 
>>> drexel_jupyter_logger.variable_logger(name_)
>>> assert isinstance(name_, str), "You did not provide your name as a string"
```
failure_message: You need to enter a name.
success_message: You entered a name. We hope it is your first and last name :).


Tests for q1 - Making a Graph (points: None)
```
>>> import drexel_jupyter_logger
>>> 
>>> drexel_jupyter_logger.variable_logger(name_)
>>> assert isinstance(name_, str), "your name is not defined"
>>> drexel_jupyter_logger.variable_logger(np.__version__)
>>> assert np.__version__ is not None, "numpy is incorrectly imported"
```
failure_message: numpy incorrectly implemented.
success_message: numpy correctly implemented.
```
>>> import drexel_jupyter_logger
>>> 
>>> drexel_jupyter_logger.variable_logger(name_)
>>> assert isinstance(name_, str), "your name is not defined"
>>> drexel_jupyter_logger.variable_logger(plt.__package__)
>>> assert plt.__package__ is not None, "plt is incorrectly imported"
```
failure_message: plt random incorrectly implemented.
success_message: plt correctly implemented.
```
>>> 
>>> drexel_jupyter_logger.variable_logger(draw_card_eq)
>>> assert callable(draw_card_eq), "draw_card_eq is not a function"
>>>
```
failure_message: draw_card_eq is not a function
success_message: draw_card_eq is callable.
```
>>> 
>>> import matplotlib.pyplot as plts
>>> 
>>> drexel_jupyter_logger.variable_logger(plot_[0])
>>> assert isinstance(plot_[0], plts.Line2D), "You did not return a plot object."
```
failure_message: You did not return a plot object.
success_message: You successfully returned a plot object.
```
>>> import drexel_jupyter_logger
>>> 
>>> drexel_jupyter_logger.variable_logger(sum(plot_[0].get_xdata()[::3]))
>>> assert np.isclose(sum(plot_[0].get_xdata()[::3]), 1.6561871562318847e-07), "The values for x_data in the plot are incorrect. "
>>> assert np.isclose(sum(plot_[0].get_ydata()[::4]), 1.2500506348082654), "The values for y_data in the plot are incorrect. "
```
failure_message: The values for your plot are incorrect.
success_message: The values for your plot are correct.
```
>>> 
>>> assert (plot_[0].get_color() == "red"
...     or plot_[0].get_color() == (1.0, 0.0, 0.0, 1.0)
...     or plot_[0].get_color() == "r"), "The line on the plot is not the correct color"
```
failure_message: The plot is not red.
success_message: The plot is red as expected.


Tests for q2 - Exploring DataTypes (points: None)
```
>>> 
>>> import drexel_jupyter_logger
>>> 
>>> drexel_jupyter_logger.variable_logger(int_)
>>> assert isinstance(int_, int)
>>> 
>>> from cryptography.fernet import Fernet
>>> 
>>> fernet = Fernet(drexel_jupyter_logger.key)
>>> out = str(fernet.decrypt(b"gAAAAABkDQ71dNmKvaPlZOj3cLT2FVfh66b09i0yV2Fl08bcFY0K_Io7ju7utXkX3cCmsb9HMsHTAHERX0X48acRp1MGHYXxSg==").decode("ascii"))
>>> 
>>> assert str(int_) == out, "int_ incorrectly implemented."
```
failure_message: int_ incorrectly implemented.
success_message: int_ correctly implemented.
```
>>> 
>>> import drexel_jupyter_logger
>>> 
>>> drexel_jupyter_logger.variable_logger(float_)
>>> assert isinstance(float_, float)
>>> 
>>> from cryptography.fernet import Fernet
>>> 
>>> fernet = Fernet(drexel_jupyter_logger.key)
>>> out = str(fernet.decrypt(b"gAAAAABkDRABl3uUfW0G0yLb5kCpVpz5XBRYUjowgJ9_c9CNAGielysMhb5cp72zDF0Gi9abZrATkUoZfUGv7UJp65Cm7WvjiQ==").decode("ascii"))
>>> 
>>> assert str(float_) == out, "float_ incorrectly implemented."
```
failure_message: float_ incorrectly implemented.
success_message: float_ correctly implemented.
```
>>> 
>>> import drexel_jupyter_logger
>>> 
>>> drexel_jupyter_logger.variable_logger(bool_)
>>> assert isinstance(bool_, bool)
>>> 
>>> from cryptography.fernet import Fernet
>>> 
>>> fernet = Fernet(drexel_jupyter_logger.key)
>>> out = str(fernet.decrypt(b"gAAAAABkDRDR42csyMOkhel0VkaN7JXse5_3OFTam1ryIyiJvEdIcfDCfocojIiGZU2PPJQm-rrfBNXLyzMQ-QJgmZi7qEdthA==").decode("ascii"))
>>> 
>>> assert str(bool_) == out, "bool_ incorrectly implemented."
```
failure_message: bool_ incorrectly implemented.
success_message: bool_ correctly implemented.
```
>>> 
>>> import drexel_jupyter_logger
>>> 
>>> drexel_jupyter_logger.variable_logger(string_)
>>> assert isinstance(string_, str)
>>> 
>>> from cryptography.fernet import Fernet
>>> 
>>> fernet = Fernet(drexel_jupyter_logger.key)
>>> out = str(fernet.decrypt(b"gAAAAABkDRECpI3kPdhirMzwLAddFoJKDXs7mXJ5AWPYJTx_gTKGYY9VR0REjaPKLCAc0MQ5UuSr4cMcBBc8zSVg2GT6LiDvLg==").decode("ascii"))
>>> 
>>> assert str(string_) == out, "string_ incorrectly implemented."
```
failure_message: string_ incorrectly implemented.
success_message: string_ correctly implemented.
```
>>> 
>>> import drexel_jupyter_logger
>>> 
>>> drexel_jupyter_logger.variable_logger(list_)
>>> assert isinstance(list_, list)
>>> 
>>> from cryptography.fernet import Fernet
>>> 
>>> fernet = Fernet(drexel_jupyter_logger.key)
>>> out = str(fernet.decrypt(b"gAAAAABkDRFLoBpFP_06PNjYcIdyOpLGFYYboKIjMjTAu936eLjoKByeaF5XTyYcvGh5PWyVVU1kFUqWj4YeqbRDfMTRpPYXGw==").decode("ascii"))
>>> 
>>> assert str(list_) == out, "list_ incorrectly implemented."
```
failure_message: list_ incorrectly implemented.
success_message: list_ correctly implemented.
```
>>> 
>>> import drexel_jupyter_logger
>>> 
>>> drexel_jupyter_logger.variable_logger(tuple_)
>>> assert isinstance(tuple_, tuple)
>>> 
>>> from cryptography.fernet import Fernet
>>> 
>>> fernet = Fernet(drexel_jupyter_logger.key)
>>> out = str(fernet.decrypt(b"gAAAAABkDRGEX2Cl-qHsEg1NFnwyZSHVLsjlc1JJWn4Lj76LT7dsEoNqlBpXQry0zKyFS2oCSzdJbYSm95SKxB96pM_HHsyuNw==").decode("ascii"))
>>> 
>>> assert str(tuple_) == out, "tuple_ incorrectly implemented."
```
failure_message: tuple_ incorrectly implemented.
success_message: tuple_ correctly implemented.
```
>>> 
>>> import drexel_jupyter_logger
>>> 
>>> drexel_jupyter_logger.variable_logger(dict_)
>>> assert isinstance(dict_, dict)
>>> 
>>> from cryptography.fernet import Fernet
>>> 
>>> fernet = Fernet(drexel_jupyter_logger.key)
>>> out = str(fernet.decrypt(b"gAAAAABkDRG1TzclExNe_t7qeJgaHYC9uWufJOfW5u7L198U2jtPxlFPQTnL9qj0MFHrNC7pIUfpz7od9FbuyQy6nuNnEFvlk2ZWGw3xRiU_wh4fVtRuV-Y=").decode("ascii"))
>>> 
>>> assert str(dict_) == out, "dict_ incorrectly implemented."
```
failure_message: dict_ incorrectly implemented.
success_message: dict_ correctly implemented.


Tests for q3-Building a Thermometer (points: None)
```
>>> 
>>> import drexel_jupyter_logger
>>> 
>>> drexel_jupyter_logger.variable_logger(name_)
>>> assert isinstance(Thermometer, type), "The Class Thermometer is not implemented."
>>> drexel_jupyter_logger.variable_logger(Thermometer)
```
failure_message: The Class Thermometer is not implemented.
success_message: The Class Thermometer is implemented.
```
>>> 
>>> import drexel_jupyter_logger
>>> 
>>> t = Thermometer(25)
>>> assert hasattr(t, "temperature_c"), "Thermometer does not have 'temperature' attribute"
>>> assert t.temperature_c == 25, "Thermometer initialization did not set attribute temperature correctly"
>>> drexel_jupyter_logger.variable_logger(t.temperature_c)
```
failure_message: The Class Thermometer initialization is implemented incorrectly.
success_message: The Class Thermometer initialization is implemented correctly.
```
>>> 
>>> t = Thermometer(25)
>>> t.set_temperature(50)
>>> assert t.temperature_c == 50.0, "set temperature does not function properly"
>>> 
>>> t = Thermometer(25)
>>> t.set_temperature("abc")
>>> out = str(t.temperature_c)
>>> assert out =="fail", "set temperature does not return the correct response when a unsuitable input is provided"
```
failure_message: set_temperature can not correctly handle non numeric inputs.
success_message: set_temperature can correctly handle non numeric inputs.
```
>>> 
>>> import drexel_jupyter_logger
>>> 
>>> t = Thermometer(0)
>>> assert t.to_fahrenheit() == 32.0, "Incorrect conversion to Fahrenheit."
>>> 
>>> t.set_temperature(20)
>>> assert t.to_fahrenheit() == 68.0, "Incorrect conversion to Fahrenheit."
>>> drexel_jupyter_logger.variable_logger(t.to_fahrenheit())
```
failure_message: Incorrect conversion to Fahrenheit.
success_message: Correct conversion to Fahrenheit.
```
>>> 
>>> import drexel_jupyter_logger
>>> 
>>> assert hasattr(water_sensor, "temperature_c"), "Thermometer does not have 'temperature' attribute"
>>> assert hasattr(water_sensor, "to_fahrenheit"), "Thermometer does not have 'to_fahrenheit' attribute"
>>> assert temp_c == 12, "Incorrect temperature for temp_c."
>>> assert temp_f == 53.6, "Incorrect temperature for temp_f."
```
failure_message: Incorrect conversion to Fahrenheit.
success_message: Correct conversion to Fahrenheit.


Tests for q4-Determining a Price for a Circuit (points: None)
```
>>> 
>>> import drexel_jupyter_logger
>>> 
>>> drexel_jupyter_logger.variable_logger(name_)
>>> assert parts_cost["10-ohm resistor"] == 0.025, "10-ohm resistor not correctly defined"
>>> assert parts_cost["100-ohm resistor"] == 0.03, "100-ohm resistor not correctly defined"
>>> assert parts_cost["1-pF capacitor"] == 0.05, "1-pF capacitor not correctly defined"
>>> assert parts_cost["10-pF capacitor"] == 0.12, "10-pF capacitor not correctly defined"
>>> drexel_jupyter_logger.variable_logger(parts_cost["10-pF capacitor"])
```
failure_message: parts_cost is incorrect. Please make sure you did not modify the dictionary.
success_message: parts_cost is correct.
```
>>> import drexel_jupyter_logger
>>> 
>>> drexel_jupyter_logger.variable_logger(name_)
>>> assert isinstance(name_, str), "your name is not defined"
>>> drexel_jupyter_logger.variable_logger(np.__version__)
>>> assert np.__version__ is not None, "numpy is incorrectly imported"
```
failure_message: numpy incorrectly implemented.
success_message: numpy correctly implemented.
```
>>> 
>>> import drexel_jupyter_logger
>>> 
>>> drexel_jupyter_logger.variable_logger(name_)
>>> assert circuit_parts["10-ohm resistor"]["quantity"] == 10, "10-ohm resistor quantity not correctly defined"
>>> assert circuit_parts["10-pF capacitor"]["quantity"] == 3, "100-ohm resistor quantity not correctly defined"
>>> assert circuit_parts["1-pF capacitor"]["quantity"] == 7, "1-pF capacitor quantity not correctly defined"
>>> assert circuit_parts["Custom Processor"]["quantity"] == 2, "Custom Processor quantity not correctly defined"
>>> assert circuit_parts["Custom Processor"]["price"] == 12, "Custom Processor price not correctly defined"
>>> drexel_jupyter_logger.variable_logger(circuit_parts["Custom Processor"]["price"])
```
failure_message: circuit_parts is incorrect. Please make sure you did not modify the dictionary.
success_message: circuit_parts is correct.
```
>>> 
>>> import drexel_jupyter_logger
>>> 
>>> drexel_jupyter_logger.variable_logger(name_)
>>> # Check function signature and name
>>> assert price_calculator.__name__ == "price_calculator", "you did not name price_calculator correctly"
>>> assert price_calculator.__code__.co_argcount == 3, "price calculator does not have the correct number of inputs"
>>> drexel_jupyter_logger.variable_logger(price_calculator.__code__.co_argcount)
```
failure_message: price_calculator call is incorrect.
success_message: price_calculator call is correct.
```
>>> 
>>> import drexel_jupyter_logger
>>> 
>>> 
>>> from cryptography.fernet import Fernet
>>> 
>>> fernet = Fernet(drexel_jupyter_logger.key)
>>> 
>>> 
>>> out = str(fernet.decrypt(b"gAAAAABkDz6Y_aio_akqMGkihfI5zKR2KkUg0yXVVHNrVV82qunlAwRjVZN2R2OYzUqkBBLBG-qd4t_2Vpx-iHxYPLMGTYQsFA==").decode("ascii"))
>>> assert str(price_calculator(parts_cost, circuit_parts, 0.2)) == out, "price_calculator  computation is incorrect."
Total cost of the order: $29.95
```
failure_message: price_calculator  computation is incorrect.
success_message: price_calculator computation is correct.
```
>>> 
>>> parts_2 = parts_cost.copy()
>>> parts_2["10-ohm resistor"] = 10
>>> 
>>> from cryptography.fernet import Fernet
>>> 
>>> fernet = Fernet(drexel_jupyter_logger.key)
>>> out = str(fernet.decrypt(b"gAAAAABkD3dtfMyEiX9AYR1pum4OZ5xSbbZ367hae3wxg5SSAcor4m5CXO6u1OthJkNgyGQUotRmCY2v18f4rjxK7Z3kbCNC3A==").decode("ascii"))
>>> assert out == str(price_calculator(parts_2, circuit_parts, 0.2)), "Your code does not use the parts_cost correctly"
Total cost of the order: $149.65
```
failure_message: Your code does not use the parts_cost correctly.
success_message: Your code correctly uses the parts_cost.
```
>>> 
>>> from cryptography.fernet import Fernet
>>> 
>>> fernet = Fernet(drexel_jupyter_logger.key)
>>> out = str(fernet.decrypt(b"gAAAAABkD3D3Z-O5XJGgN67FRSetS5oiOetUTzWxQFbw_OTKnrjBEhnelTKcH0d22z4AneFL8pYrtOgSg0Q2qYc0A4dI-VdzzcV6pXrv5aKw-FAVmhHMZ7U=").decode("ascii"))
>>> 
>>> 
>>> assert print_total(price_calculator(parts_cost, circuit_parts, profit)) == out, "The string being printed is not correct"
Total cost of the order: $40.68
```
failure_message: print message is incorrect.
success_message: print message is correct.


Tests for q5-Plotting and Fitting (points: None)
```
>>> 
>>> import drexel_jupyter_logger
>>> 
>>> drexel_jupyter_logger.variable_logger(name_)
>>> assert "numpy" in np.__name__, "numpy not imported correctly"
>>> assert "pyplot" in plt.__name__, "matplotlib.pyplot not imported correctly"
>>> assert "curve_fit" in dir(), "scipy.optimize.curve_fit not imported correctly"
>>> drexel_jupyter_logger.variable_logger(plt.__name__)
```
failure_message: Your input statements are incorrect.
success_message: Your input statements are correct.
```
>>> 
>>> import drexel_jupyter_logger
>>> 
>>> 
>>> # test values
>>> x = np.array([0, 1, 2])
>>> a_, b_, c_ = 1, 1, 1
>>> expected_output = np.array([0.60653066, 1.0, 0.60653066])
>>> 
>>> # check if the function returns the expected output for test values
>>> assert np.allclose(gaussian(x, a_, b_, c_), expected_output), "gaussian function is not implemented correctly"
>>> 
>>> x_test = np.array([0, 1, 2])
>>> a_test, b_test, c_test = 1, 1, 1
>>> expected_output = np.array([0.60653066, 1.0, 0.60653066])
>>> assert np.allclose(gaussian(x_test, a_test, b_test, c_test), expected_output), "gaussian function is not implemented correctly"
>>> drexel_jupyter_logger.variable_logger(gaussian(x, a_, b_, c_))
```
failure_message: gaussian function is not implemented correctly
success_message: gaussian function is implemented correctly
```
>>> 
>>> import drexel_jupyter_logger
>>> 
>>> assert np.allclose(x_data[0], -20), "x_data is not defined correctly"
>>> assert np.allclose(x_data[-1], 20), "x_data is not defined correctly"
>>> assert len(x_data) == 100, "x_data is not defined correctly"
>>> assert np.isclose(sum(x_data[3::3]), 19.99999999999997), "x_data is not defined correctly"
>>> 
>>> drexel_jupyter_logger.variable_logger(len(x_data))
```
failure_message: x_data is not defined correctly.
success_message: x_data is defined correctly.
```
>>> 
>>> import drexel_jupyter_logger
>>> 
>>> # Checks y_true
>>> from cryptography.fernet import Fernet
>>> 
>>> fernet = Fernet(drexel_jupyter_logger.key)
>>> out = str(
...     fernet.decrypt(
...         b"gAAAAABkD-NzYrMglFaqNyS-zlkoQo-gwxc1YCD39uP-uNVNborNwKBB7YKbPEpayy9rq7hUGxJhZWrWmXBDQ14fxazkpMi5v-8ArClYcJqu2ZJvUs2LKzPwphZZ63pfZBxs8IK0DIw-EnU5mRoiaMT3YoTAD7rSGUrnDdqM44YGV3b_Gyt7mYE0n_obgjbXngI0DbuJzExiXeA4KQNpSsKJD6x870owlz7wNXADMhSrU03suFLlUx1pgGwJsg-7lQhUocs64ZahOJvSfABgqu_Ml_SDkaudvTyHTFCPy-bnq3xn-Cc4d9jmsdOxUmqMvl3xLZw1ALt6YrMBdlNZh0hBgjrT4iVP9tySJBG8dnov7TLz9_PuWRGRulKG-tYZU9XJ8uuFC09mXpDLE49gdlCDp_jYbvYfWX6rPlgBd7pXyQ2YWAKajlRzuJ1Q7TrXSVm_Xbj2eF48_iehGgz2cuzg0YARRKXvrg3JFdRn-jRhN9Pop4SoJVEcIbh6p6AmU0u2rUTFDN5UazGdLSWav8g2xQyHCm8_2i9il6lVF3_H4E5gdycNi6eb4hUV0i7XY1rUzfulJYmJzWzQk3KgbR54In9AfLWvgHY0IcBduStXBmpMfwIG7NMbbSPhElQ0D6gXjA3F7zpkahXZ3S0ihF0bfHvY32pIfRVORtMT811bYif8OPTJH2i3YcHiaS7ElIUJXkb4dnJTOOVCIB08EV0ONZgZHwtpsFCCwCFvRTouorcfz1SXs9a0vtjtau3CgBI3cl0yAGGuiazfipZAEsbB68QedFffR_QRl0ZBN_IfGIMC8ksxUw3R9J46tUET_doctLBvd7LFbDTfkECNoKCh-E5x-POLyjeWgscHe64WSRNkYgKaKTWP70Y9OdYg2CpCZseE55ovOqgNjTVCMlRhuQWO9zuBSLAU3NOfL-ms4n9AWrx75Wsg9gxkIGVm36HZ-yKO_wJDcrRRQqtYaS0noeqEdo3-Gmj4ldMgPOATkNXKLoFJ7P6FfnNYUiY7mPykn1_1tqlBEhVvSXrAteRNr6E_8s6lgFr4l2kDU8s8_XSgKu_wazGqw45WVMiO5Zgsh2SXh4VnnUE0LFubdThZsKCplGr_uJK2IFc0jiWlSid8rE2GiLU69Y0mVkwod2PQOfIfJRzeNkAA9raRbs5CkFO8XFVgLGjKWGIJw_8zKncKHZ-LgK0PVFJ2zMBMuN8fXofonS3v7Zm85gj16MGtfwIOVhvODG4BrOkZvI43TCM5pPRq1-5nqyJfJ6gpC0SWWwMrEF7YSZ64YIA5YmVlaquaB5Fam4p5O3EN8m83aQ4AcCBcRZAaimCtPIrIhK7SS394wkfQLFb8q_WUJe5RBHm3PoVsbjac3Vlu7rQTVByxSf9KdFNdtKTXJsMzLL0743ZMBmm4RA5UlQWE15AFeIyiSXP8B0TgFKe9J5g8W-eT8KWyqRw-x00_7lF0e-4tcB1B6CLOuk4kpkZk0WkucN0S1IzZRZsjXcOwmsoNDqFdpx7NrO08ZgmXqkSqwLCo99q0C6gpk2oX3-BBw9cDOgvFIpYE4A=="
...     ).decode("ascii"))
>>> assert str(y_true) == out, "y_true is incorrect"
>>> 
>>> drexel_jupyter_logger.variable_logger(len(y_true))
```
failure_message: y_true is incorrect.
success_message: y_true is correct.
```
>>> 
>>> import drexel_jupyter_logger
>>> 
>>> # Checks y_noisy
>>> from cryptography.fernet import Fernet
>>> 
>>> fernet = Fernet(drexel_jupyter_logger.key)
>>> out = str(fernet.decrypt(b"gAAAAABkD-UVvxc6wECqC_U9taa6Art0EtaZg2i_WB0sp4FHCf3Ggyed8uUrljAkCPxg-AwYxN2jnqaPVgMh57F1OX1ftyjzR1SfhzBk5UAkxPaLIFHteLZY_7oK67GesWvX-jys-fhfKLktB_4ba54nXaVQP3K2xb60izc-42Y4VvddXu9Iggo5dy7VcJP8-6Ix_un8KqJo2J4CKYXFni2El-O9zrvvc1EcNYgvGdpzxFFeUHEHAP0AEqTp-Ow5D7T0x2iGcEPXtDUcmbgMG_t-HAQRr8ky9DNyuKx9k13pWT85O3jpLAtExhBYk2yXtzS9iU4pQEYGpYhdB2HQosUQcV5rXdgsC6pGcWlWhTab2uG32_Lt8-Y0jlb4VgALBAtxJGXuevGpJ7Pt1fzt7pjHikCVHcMbmkhLaWGzZq1uf6d6mbA-lxo1wGif7-ryFB4TTyze2HoxcGGgLI7vtuHgduBSd397OTTuS9PMMnR7oIXBZprj_NJrSmI79N4mMDlkWz-baZikhLw6LbThxinvd4HX6aQPVm3bPlV9M5Yc3lSqr2YQimUKqONnZGKxzcQx4TxfeRjdv8mS49Qey1KRU5zUlBS11NX-ScrqIFDAwFBsq6OpMSqd80rZquYE5E-IsHr0g4McffuDal20JZqXtP_kYMaT3Iv_KCa__9TKmpMHKJpQ44qhvzuBkqA59npHU05s5gbr5OpPwug1W_H7WTgjRYD-IaDU9cxxKWs9erfEkcCMh7peMvnBVwu2IY_W1u7ocNxXiMR2BJX6FbGO2P12t7D8vz3lyCJFHhkYggO06v9UF0L3-eX186gXNdegYnML3kVevZbOkRHFxvU2hoZ6dzLW8Z8SgaO9H1-Pac5Ima1Kh36EF1GH-93QGnSf6w2feXG9yByfSSbRharfHT6-79nYgl3EgGL0wo8imUaN8uwRhkxma8myGCAsNbkPKPiltTXfyv8yKEjc-qC5UivinYtFYYkhiDqIQvQDB4YURk0p4MxesM1rOgq6aML_jc7haRUcUes4ynGERvVt_Nc3bV03--BIzTZx4D2q3OMEUt25CyTU9oUr6r4C_JsrxhtK2d5U9qrLY6b-W-F-lBI3ZrRh_HcwHR9Hx7YTA7vRcXaOyYB-fUWjwXa2Gbxj1s9iEpzWf8lEisQk40FUhc2jSmtiZEn6MO3dO_YxJ3s9TE1Tmq9aK7gmb3zMbZn2gQyuHxzuvnUeJBi-DelNLvQXP7C3GYyMjA6TuajEzNY6VSVE7sKgQyw_vjE0QteWtkUpODlAbpT42M6laN3ahNeT7l2WE0D-wvjEHzflGOF7udat6QccodlHUge4DAJSS4vaTWs5J-aoSCmPNrYN65W3CIokMhxXS0VZ0n26ObdU384YVKdOTtUn796D7Mdk0ll0ImLLs_1RcihuZPkmEhE2rnWUINRBOLAyoQz6RrNGe-tC-uA3XNsxYbDQ9ReBPlU_R_BbkkgEzKtMV7juGtddtKDF-1lzhoahFX33CNMdWcuq189G87JsEED_tdmSloByFx3sEExB0uuj0--pnbH-CV1674xJ8x2WarCevNkOts1Aa0tDyE9v9v8NmX6cIKHsT8rkIhStCE1_Xknn2uMGlqNPp_5k30P5_mKeth6tti6uH-MMNO_qQ8V82Ki7X5u3HOmPgAa3_uJGD-JVvREQ8uNw8SVFpCmO-X207mkqYEFXO_Q=").decode("ascii"))
>>> assert str(y_noisy) == out, "y_noisy is incorrect"
>>> 
>>> drexel_jupyter_logger.variable_logger(len(y_noisy))
```
failure_message: y_noisy is not implmemented correctly
success_message: y_noisy is implmemented correctly
```
>>> 
>>> import drexel_jupyter_logger
>>> 
>>> # Checks popt
>>> from cryptography.fernet import Fernet
>>> 
>>> fernet = Fernet(drexel_jupyter_logger.key)
>>> out = fernet.decrypt(b"gAAAAABkD-VjEnvkOx1q5hIc6OiMyPE7iayGWAtZR_Kry0mvrQbCr0ffV9Lrck5j673YFnlWxRjo9qR44gKpYzUOzGhmCBpXKZ2Dec1hHWXhY4szHZZuirxAWmfMc66dj5dpA1bzjYFG").decode("ascii")
>>> chars_to_remove = ["[", "]","\n"]
>>> for char in chars_to_remove:
...     out = out.replace(char, "")
>>> 
>>> out = [float(x) for x in out.split(" ")]
>>> 
>>> assert np.allclose(popt, out), "popt is not returning the correct values"
>>> 
>>> drexel_jupyter_logger.variable_logger(popt)
```
failure_message: popt is not returning the correct values
success_message: popt is returning the correct values
```
>>> 
>>> import drexel_jupyter_logger
>>> 
>>> # Checks y_fit
>>> from cryptography.fernet import Fernet
>>> 
>>> fernet = Fernet(drexel_jupyter_logger.key)
>>> out = fernet.decrypt( b"gAAAAABkD-WzDT72FrJF_4PrNsTIhnppKucv8pAjRnotXNrslY7YlSdDyqgiXabo32_bwcFXYkokXar7Vl-Bl1cN4lD70-PDCbfDPsxVDcGs25T4Z8IV-i6P7SAwaZ-k00NmVsJ4sF62ClW9ssomSd0jgYtbgy4J5sF9_KO-6s0ke6_Pfi1mrx8c90J5ZYsAwB9NKm45dZ-IWfjH955EkoxoPQcdqBQ16RG1J0k5DqGaP6hpRpR4mRANk50AFcgkX9PLqnKr3Z3AH1s_ugMEvza1GRlj1aVmaeF0JcHsVw1MxpZLLhPK0ikm8TcyX7cy6O8wi3xvjk5By8SFT96a0pqK5ObXGIuFAN6XlBH56yic8ttcRBODzaRiH_97wGvi8eUP0NNGOlMZFb_mPc7E9URWZCnJ0Krr57iOj5ALWDPqtMysnRGg4VCI9l6ZAFv0dP_HqbZT7U7TK_dFVAWNWv5MPeyg8rcGInJ2ygH1o1G8jRJRSbEcHQj3-ZU2K7DCFqc6xg2EO_31Dz6VLWnEBLVW2OUIth16vPCcjbmJgGuW0mgRulbogMNLFEXpGKJwi8LimBL_gZ0pee4vd8nZDKaRR8pSoLXsHyTHf-O8vXHzji4RiUF7VSReSrbpai1xoC7uyztjO1Vjce1bMf1pP7e3JNIzpJ7sfAlP0PPiERdG58Xe_2Qx_oWFMtQeDt-ViS4AWK76KlWgsiX0cFqWs3Dj_KMs1hmtywbEZtLkUjEUH6rclgw482Tp44fFQcUqBOOD92DJkGjmEr-FcVJtlu8R71hafSMBnOJ_cEUvZL-P5ljNUoUoM1I7REUK9hs-hiSYEyXbyCfv5dRSoXHOdNv0nDKoEcVWPYZPP3E74Qafmm47NHUNbcG3pMGXXMsOea7l5SnL_gfnYjNcvUIDEzZDXen2moiSP0P5EIeFTSItTdjSnM-TAMmDZlFSF3XtXD8vagpPb_-QdfPNOkkx1MQNjVXcKqDVPUwyuDkrqm0QFp2XrRW8Lwb7v5mENaD3HfZMlnpuBYAv9sXDtkb7Km30U9NgYExvScKEnBfbd61peROkRYXJiTA8uvx3xgWAeu47lCYNJ0jKomOQlJOVXH86TnPEjMKDQ9UUcFTJGcuQLaNfSPdSPqkf1gabg3e6VW4b--PSkhpss9OboUp1UOEo0KFnlj4Tl-zu8rwuWEUwlBUDYIyuKwWBYD_DpC62kBfCcuMuaM1xvbbMoedQLTm0OA0PIR_dP5nrW4IkGtnpbz43rqBMFCcGZ2un8KbCRT1mh8dB5nNdrtxl3dtBEB_nsl-M8gIqCruIfNHkoPx77SkXHT3zNnI66oYqgBpMCrTjShWB-xNtlRrKHqL0q9RpgPF9SJphcLBQ07eqV4s7TrJtfrhw41WuzLaCR7KYb_vsMj78KNrZmuaTDgGGu4lediwCYGOFA3I-jIxSj2jwYfDcrGKW9QUbDK5RTZ0_rEt6SwA5r7NpadXtv-RsGBycn-_vTNPQ_5PGIYGyY3GkQgVjiq0nBX1xI39PqmMKc0eyGKAH6pu2uhCw_FAe3GSSHzMke-zzIA==").decode("ascii")
>>> chars_to_remove = ["[", "]"," \n"]
>>> for char in chars_to_remove:
...     out = out.replace(char, "")
...     
>>> out = out.replace("  ", " ")
>>> 
>>> out = [float(x) for x in out.split(" ")]
>>> 
>>> assert np.allclose(y_fit, out), "y_fit is returning incorrect values"
>>> 
>>> drexel_jupyter_logger.variable_logger(popt)
```
failure_message: y_fit is returning incorrect values
success_message: y_fit is returning correct values
```
>>> 
>>> # Check the xlabel
>>> assert plot_1.axes.get_xlabel() == "x", "xlabel is incorrect"
>>> 
>>> # Check the ylabel
>>> assert plot_1.axes.get_ylabel() == "y", "ylabel is incorrect"
```
failure_message: The plot labels are incorrect.
success_message: The plot labels are correct.
```
>>> 
>>> # Check that the label of the data points is correct
>>> assert plot_1.get_label() == "Noisy Data", "plot_1 is not labeled correctly"
>>> 
>>> # Check that the label of the data points is correct
>>> assert plot_2[0].get_label() == "Fitted Gaussian", "plot_2 is not labeled correctly"
```
failure_message: the plots are not labeled correctly
success_message: the plots are labeled correctly
```
>>> 
>>> # Check that the color of the fitted Gaussian line is correct
>>> assert plot_2[0].get_color() == (0.0, 0.5, 0.0, 1), "gaussian fit is not the correct color"
```
failure_message: gaussian fit is not the correct color
success_message: gaussian fit is the correct color
```
>>> 
>>> # Check that the legend is displayed
>>> assert plot_1.axes.get_legend() is not None, "The legend is not correct on the plot"
```
failure_message: The legend is not correct on the plot
success_message: The legend is correct on the plot
```
>>> 
>>> assert all(plot_2[0].get_xdata() == x_data), "the x data for plot_2 is not correct"
>>> assert all(plot_2[0].get_ydata() == y_fit), "the y data for plot_2 is not correct"
```
failure_message: plot_2 does not have the correct values.
success_message: plot_2 is correct.
```
>>> 
>>> assert np.array_equal(plot_1.get_offsets()[:, 1], y_noisy), "the y data for plot_1 is not correct"
>>> assert np.array_equal(plot_1.get_offsets()[:, 0], x_data), "the x data for plot_1 is not correct"
>>>
```
failure_message: plot_1 does not have the correct values.
success_message: plot_1 is correct.
